[1613561579] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 401 Unauthorized - 691 bytes in 0.000 second response time
[1613561639] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 401 Unauthorized - 691 bytes in 0.002 second response time
[1613561692] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;1;connect to address 127.0.0.1 and port 22: Connection refused
[1613561699] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;3;HTTP WARNING: HTTP/1.1 401 Unauthorized - 691 bytes in 0.001 second response time
[1613561729] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;1;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613561752] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;2;connect to address 127.0.0.1 and port 22: Connection refused
[1613561759] SERVICE ALERT: localhost;HTTP;WARNING;HARD;4;HTTP WARNING: HTTP/1.1 401 Unauthorized - 691 bytes in 0.001 second response time
[1613561789] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;2;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613561812] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;3;connect to address 127.0.0.1 and port 22: Connection refused
[1613561849] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;3;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613561872] SERVICE ALERT: localhost;SSH;CRITICAL;HARD;4;connect to address 127.0.0.1 and port 22: Connection refused
[1613561909] SERVICE ALERT: localhost;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613567504] SERVICE ALERT: localhost;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 16.05, 6.36, 3.15
[1613567564] SERVICE ALERT: localhost;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 15.04, 7.72, 3.82
[1613567624] SERVICE ALERT: localhost;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 13.31, 8.89, 4.49
[1613567684] SERVICE ALERT: localhost;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 5.44, 7.47, 4.27
[1613567984] SERVICE ALERT: localhost;Current Load;WARNING;HARD;4;WARNING - load average: 0.70, 3.10, 3.24
[1613568284] SERVICE ALERT: localhost;Current Load;OK;HARD;4;OK - load average: 1.10, 1.64, 2.55
[1613573714] SERVICE ALERT: viewnext112;HTTP;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 401 Unauthorized - 691 bytes in 0.000 second response time
[1613573783] SERVICE ALERT: viewnext112;Jenkins;CRITICAL;SOFT;1;connect to address 127.0.0.1 and port 9080: Connection refused
[1613573834] SERVICE ALERT: viewnext112;HTTP;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 401 Unauthorized - 691 bytes in 0.001 second response time
[1613573853] SERVICE ALERT: viewnext112;ssh;CRITICAL;SOFT;1;connect to address 127.0.0.1 and port 22: Connection refused
[1613573903] SERVICE ALERT: viewnext112;Jenkins;CRITICAL;SOFT;2;connect to address 127.0.0.1 and port 9080: Connection refused
[1613573954] SERVICE ALERT: viewnext112;HTTP;WARNING;HARD;3;HTTP WARNING: HTTP/1.1 401 Unauthorized - 691 bytes in 0.001 second response time
[1613573973] SERVICE ALERT: viewnext112;ssh;CRITICAL;SOFT;2;connect to address 127.0.0.1 and port 22: Connection refused
[1613574023] SERVICE ALERT: viewnext112;Jenkins;CRITICAL;HARD;3;connect to address 127.0.0.1 and port 9080: Connection refused
[1613574096] SERVICE ALERT: viewnext112;ssh;CRITICAL;HARD;3;connect to address 10.250.12.2 and port 22: Host is unreachable
[1613574167] HOST ALERT: viewnext112;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1613574241] HOST ALERT: viewnext112;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1613574315] HOST ALERT: viewnext112;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1613574379] HOST ALERT: viewnext112;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1613574555] SERVICE ALERT: viewnext112;HTTP;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 687 bytes in 0.231 second response time
[1613574626] SERVICE ALERT: viewnext112;Jenkins;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 55950 bytes in 2.789 second response time
[1613574696] SERVICE ALERT: viewnext112;ssh;OK;HARD;3;SSH OK - OpenSSH_7.6p1 Ubuntu-4ubuntu0.3 (protocol 2.0)
[1613574849] SERVICE ALERT: viewnext112;Gitlab;CRITICAL;SOFT;1;connect to address 10.250.0.2 and port 8929: Connection refused
[1613574889] SERVICE ALERT: viewnext112;Gitlab;CRITICAL;SOFT;2;connect to address 10.250.0.2 and port 8929: Connection refused
[1613574890] SERVICE ALERT: viewnext112;Jenkins;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 401 Unauthorized - 782 bytes in 0.117 second response time
[1613575009] SERVICE ALERT: viewnext112;Gitlab;CRITICAL;HARD;3;connect to address 10.250.0.2 and port 8929: Connection refused
[1613575010] SERVICE ALERT: viewnext112;Jenkins;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 401 Unauthorized - 782 bytes in 0.107 second response time
[1613575188] SERVICE ALERT: viewnext112;Jenkins;WARNING;HARD;3;HTTP WARNING: HTTP/1.1 401 Unauthorized - 782 bytes in 0.109 second response time
[1613576079] SERVICE ALERT: viewnext112;HTTP;CRITICAL;SOFT;1;connect to address 10.250.12.2 and port 80: Host is unreachable
[1613576079] SERVICE FLAPPING ALERT: viewnext112;Jenkins;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1613576079] SERVICE ALERT: viewnext112;Jenkins;CRITICAL;HARD;3;connect to address 10.250.12.2 and port 9080: Host is unreachable
[1613576079] SERVICE ALERT: viewnext112;ssh;CRITICAL;SOFT;1;connect to address 10.250.12.2 and port 22: Host is unreachable
[1613576079] SERVICE ALERT: viewnext112;Gitlab;CRITICAL;SOFT;1;connect to address 10.250.12.2 and port 8929: Host is unreachable
[1613576079] SERVICE ALERT: viewnext112;ssh;CRITICAL;SOFT;2;connect to address 10.250.12.2 and port 22: Host is unreachable
[1613576079] SERVICE ALERT: viewnext112;HTTP;CRITICAL;SOFT;2;connect to address 10.250.12.2 and port 80: Host is unreachable
[1613576093] HOST ALERT: viewnext112;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1613576167] HOST ALERT: viewnext112;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1613576202] SERVICE ALERT: viewnext112;HTTP;CRITICAL;HARD;3;connect to address 10.250.12.2 and port 80: Host is unreachable
[1613576202] SERVICE ALERT: viewnext112;ssh;CRITICAL;HARD;3;connect to address 10.250.12.2 and port 22: Host is unreachable
[1613576202] SERVICE ALERT: viewnext112;Gitlab;CRITICAL;HARD;3;connect to address 10.250.12.2 and port 8929: Host is unreachable
[1613576241] HOST ALERT: viewnext112;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1613576252] SERVICE FLAPPING ALERT: viewnext112;Jenkins;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1613576276] SERVICE ALERT: viewnext112;Jenkins;WARNING;HARD;3;HTTP WARNING: HTTP/1.1 403 Forbidden - 1531 bytes in 0.010 second response time
[1613576276] SERVICE FLAPPING ALERT: viewnext112;HTTP;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1613576276] SERVICE ALERT: viewnext112;HTTP;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 644 bytes in 0.004 second response time
[1613576276] SERVICE FLAPPING ALERT: viewnext112;ssh;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1613576276] SERVICE ALERT: viewnext112;ssh;OK;HARD;3;SSH OK - OpenSSH_7.6p1 Ubuntu-4ubuntu0.3 (protocol 2.0)
[1613576280] HOST ALERT: viewnext112;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1613576406] SERVICE FLAPPING ALERT: viewnext112;HTTP;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1613576406] SERVICE FLAPPING ALERT: viewnext112;Jenkins;STARTED; Service appears to have started flapping (29.0% change >= 20.0% threshold)
[1613576406] SERVICE FLAPPING ALERT: viewnext112;ssh;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1613576521] SERVICE FLAPPING ALERT: viewnext112;HTTP;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1613576521] SERVICE FLAPPING ALERT: viewnext112;Jenkins;STARTED; Service appears to have started flapping (28.4% change >= 20.0% threshold)
[1613576521] SERVICE FLAPPING ALERT: viewnext112;ssh;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1613576555] SERVICE ALERT: viewnext112;Jenkins;CRITICAL;HARD;3;connect to address 10.250.12.2 and port 9080: Host is unreachable
[1613576555] SERVICE ALERT: viewnext112;HTTP;CRITICAL;SOFT;1;connect to address 10.250.12.2 and port 80: Host is unreachable
[1613576555] SERVICE ALERT: viewnext112;ssh;CRITICAL;SOFT;1;connect to address 10.250.12.2 and port 22: Host is unreachable
[1613576569] HOST ALERT: viewnext112;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1613576627] SERVICE FLAPPING ALERT: viewnext112;Jenkins;STARTED; Service appears to have started flapping (33.9% change >= 20.0% threshold)
[1613576633] HOST ALERT: viewnext112;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1613576642] SERVICE FLAPPING ALERT: viewnext112;ssh;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1613576642] SERVICE ALERT: viewnext112;ssh;OK;SOFT;2;SSH OK - OpenSSH_7.6p1 Ubuntu-4ubuntu0.3 (protocol 2.0)
[1613576642] SERVICE FLAPPING ALERT: viewnext112;HTTP;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1613576642] SERVICE ALERT: viewnext112;HTTP;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 687 bytes in 0.001 second response time
[1613576642] SERVICE ALERT: viewnext112;Gitlab;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 1030 bytes in 0.256 second response time
[1613576643] SERVICE ALERT: viewnext112;Jenkins;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 53497 bytes in 0.740 second response time
[1613580586] SERVICE ALERT: localhost;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 11.15, 5.92, 3.19
[1613580646] SERVICE ALERT: localhost;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 12.99, 7.19, 3.79
[1613580706] SERVICE ALERT: localhost;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 14.37, 8.88, 4.60
[1613580767] SERVICE ALERT: localhost;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 12.64, 9.34, 5.03